```python
# Initialize Otter
import otter
grader = otter.Notebook("COGS118B_D1.ipynb")
```

# Discussion Notebook Week 1

In this notebook, we are going to do some basic reviews about python, numpy, scipy, along with some basic review with probability theory and estimation.

```python
# According to https://github.com/jmshea/jupyterquiz/issues/20
# Restart the kernel if needed
!pip install -q jupyterquiz==2.7.0a1
from jupyterquiz import display_quiz
```

## Different Types of Errors

```python
display_quiz("data/prob.json")
```

## Randon Variables and Probability Events

```python
display_quiz("data/rv.json")
```

## Expected Values of Discrete / Continous R.V.

```python
display_quiz("data/EX.json")
```

## Numpy

#### Import Numpy

```python
import numpy as np
```

#### Numpy Array Creation

Create an one-dimensional array of $\begin{bmatrix}1,2,3\end{bmatrix}$ and a two dimensional array of $\begin{bmatrix}1 \\ 2 \\ 3\end{bmatrix}$

```python
array_1d = ...

print("1D array:", array_1d, "Shape:", array_1d.shape)

array_2d = ...

print("2D Array:")
print(array_2d)
print("Shape:", array_2d.shape)
```

#### Array Indexing 

```python
# Given a 3x3 matrix
print("Matrix")
m = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]]) # Create a 3x3 array.
print(m)
```

```python
# Extract the 2nd row from the Matrix m
second_row = ...
print(second_row)
```

```python
# Extract the 3rd column from the Matrix m
third_column = ...
print(third_column)
```

```python
# Extract the element in the 1st row and 3rd column from the Matrix m
ele = ...
print(ele)
```

#### Modify an Array
Given the matrix `m`, modify the 1st row of the following 3x3 matrix to [100, 100, 100] inplace.

```python
m = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]]) # Create a 3x3 array.
print('Before Modification:')
print(m)
```

```python
...

print('After Modification:')
print(m)
```

```python
# Change the last column of the following 3x3 matrix to [100, 100, 100]
m = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]]) # Create a 3x3 array.
print('Before Modification:')
print(m)
```

```python
# Solution
...
print('After Modification:')
print(m)
```

#### Math Operations

```python
a = np.array([[1, 2, 3], [4, 5, 6]], dtype=np.float64)                                      
```

given the matrix `a`, multiply each element by 3.

```python
multiply_3 = ...
print(multiply_3)
```

The following is a list of availabe matrix operations when you have two matrices.

```python
b = np.array([[1, 1, 1], [2, 2, 2]], dtype=np.float64)

print(a + b)                                            # Elementwise sum
print(a - b)                                            # Elementwise difference
print(a * b)                                            # Elementwise product
print(a / b)                                            # Elementwise division
print(a == b)                                           # Elementwise comparison
```

```python
m = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])

# Sum of all array elements
print(np.sum(m))
# Sum of each column
print(np.sum(m, axis=0))
# Sum of each row
print(np.sum(m, axis=1)) 
```

#### Matrix Operations

```python
a = np.array([[1, 2], [3, 4]])
b = np.array([[2, 1], [3, 1]])
# Multiply matrix a with matrix b
ab = ...
print(ab)

# Multiply matrix a with a vector x
x = np.array([4, 2])  
ax = ...
print(ax)
```

#### Load Data with Numpy

```python
# The Iris dataset(with no labels) is stored in 'data/iris.txt'
# Load the dataset with Numpy
iris_dataset = ...
```

```python
# Each row represents features of sample
# Each column represents one type of feature
print(iris_dataset[:10])
```

## Matplotlib

```python
# Import matplotlib
import matplotlib.pyplot as plt
```

Plot a sine function over the domain of $[-2\pi, 2\pi]$

```python
...
```

```python
grader.check("q6")
```

## End of D1 :-)

Tests for q6 (points: None)
```
>>> assert np
>>> assert plt
```
